---

# Assignment 3 Task 2 (3 points)

You are given the following dictionary of sets:

```python
prog_skills = {
    "Anna": {"C#", "Java"},
    "Bob": {"Python", "C++", "JavaScript"},
    "Carla": {"Java", "JavaScript", "Python"},
}
```

Using this dictionary, write a function that returns the answer to each of the questions below. The returned string should follow the format of the expected output below.

### Question 1: What language(s) can a person code in?

For Anna, expected output is:
``Anna can code in {'Python', 'Java'}.``

(Order of items in a set may vary, and that's just fine.)

<!-- grade_id: cell-c8c1768e7f527c52, points: 0, grade: false, solution: true -->
```python

def question1(dictionary, person):
    for name in dictionary:
        if name==person:
            return name+" can code in "+str(dictionary[name])+"."

```

### Question 2: How many language(s) can a person code in?

For Bob, expected output is:
``Bob codes in 3 languages.``

<!-- grade_id: cell-e67bd891d067ff8a, points: 0, grade: false, solution: true -->
```python
def question2(dictionary, person):
    count=0
    for name in dictionary:
        if name==person:
            length=dictionary[name]
            return name+" codes in "+str(len(length))+" languages."
```

### Question 3: What language(s) do both person 1 and person 2 code in?

For Anna and Carla, expected output is:
``They both code in  {'Java'}.``

For Bob and Carla, expected output is:
``They both code in {'Java', 'JavaScript'}.``

<!-- grade_id: cell-628b0844864613bc, points: 0, grade: false, solution: true -->
```python
def question3(dictionary, person1, person2):
    firstAttr=dictionary[person1]
    secAttr  =dictionary[person2]
    result=firstAttr & secAttr
    return "They both code in "+str(result)+"."
```

# Tests

The cell below contains unit tests for your solution.

Run the below cell (after running all the above cells with your code).

<!-- grade_id: cell-b8fd170e6b42cf68, points: 1, grade: true, solution: false -->
```python
assert question1(prog_skills, "Anna") == "Anna can code in {'C#', 'Java'}." or \
question1(prog_skills, "Anna") == "Anna can code in {'Java', 'C#'}."
```

<!-- grade_id: cell-7aae16280978cf72, points: 1, grade: true, solution: false -->
```python
assert question2(prog_skills, "Bob") == "Bob codes in 3 languages."
```

<!-- grade_id: cell-53305c7f6d30921e, points: 1, grade: true, solution: false -->
```python
assert question3(prog_skills, "Anna", "Carla") == "They both code in {'Java'}."

assert question3(prog_skills, "Bob", "Carla") == "They both code in {'JavaScript', 'Python'}." or \
question3(prog_skills, "Bob", "Carla") == "They both code in {'Python', 'JavaScript'}."
```

```python

```