---

# Assignment 5 Task 1 (10 points)

You created the ``mytransform`` module last week. This week, let's turns that into a class, with a little extension. Please fulfill the following requirements for the class:

- The class shall be called ``MyTransform``.

- There are 4 private attributes: ``_data_original``, ``_data_flattened``, ``_row``, ``_col``. Except ``_data_flattened``, each of them shall be accessible (i.e. both get and set) publicly via the ``@property`` approach.

- There should be **NO** possibility to get and set ``_data_flattened``. (In what way does this design decision make sense?)

- The ``__init__()`` method does not take any parameters (except self, of course), but defines the default values of ``_row = 3`` and ``_col = 4``.

- The ``flatten()`` method turns ``_data_original`` into a 1-D list and stores that in ``_data_flattened``. No return value.
  
- The ``reshape()`` method works on ``_data_flattened`` and returns it in the new shape.

- The ``size()`` method returns the number of elements in ``_data_original`` or ``_data_flattened``. (They have the same size anyway.)

- The ``__repr__()`` method returns an appropriate string describing the instance. Use ``__class__`` to include the class name into the string.

A typical flow of using the class looks like this:
- Create an "empty" instance with no data
- Set the ``data_original``, which is any rectangular 2-D matrix. Rectangular means that every row has the same number of columns.
- Call ``flatten()``
- Inspect the flattened data via ``data_flattened`` and ``size()``
- Change the default ``row``, and ``col`` if needed
- Call ``reshape()``
- Change ``row``, and ``col`` again
- Call ``reshape()`` again

Calling external library / API to flatten / reshape leads to zero points! You MUST implement your own methods.

Write your entire class in the single cell below.

<!-- grade_id: cell-28047faa3ebe06b8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

class MyTransform:
    def __init__(self):
        # Initialize private attributes with default values
        self._data_original = []
        self._data_flattened = []
        self._row = 3
        self._col = 4

    @property
    def data_original(self):
        # Getter for _data_original
        return self._data_original

    @data_original.setter
    def data_original(self, value):
        # Setter for _data_original
        self._data_original = value

    @property
    def row(self):
        # Getter for _row
        return self._row

    @row.setter
    def row(self, value):
        # Setter for _row
        self._row = value

    @property
    def col(self):
        # Getter for _col
        return self._col

    @col.setter
    def col(self, value):
        # Setter for _col
        self._col = value
        
    def reshape(self):
        if len(self._data_flattened) != self._row * self._col:   # Check if the flattened data can be reshaped to the specified dimensions
            raise ValueError("The total number of elements does not match the desired shape.")
        reshaped_list = []
        for i in range(self._row):
            start_index = i * self._col  # Calculate start and end index for current slice
            end_index = start_index + self._col
            reshaped_list.append(self._data_flattened[start_index:end_index])  # Append slice to result list
        
        return reshaped_list

    def size(self):
        return len(self._data_flattened) # Return the number of elements in _data_flattened

    def flatten(self):
        self._data_flattened = []
        for sublist in self._data_original:
            for item in sublist:
                self._data_flattened.append(item)

    def __repr__(self):
         # Return a string representation of the instance
        return f"{self.__class__.__name__}(data_original={self._data_original}, data_flattened={self._data_flattened}, row={self._row}, col={self._col})" 


```

<!-- grade_id: cell-28047faa3ebe06b8, points: 0, grade: false, solution: true -->
```python

my_transform = MyTransform()
my_transform._data_original = [
    [0,1,2,3],
    [4,5,6,7],
    [8,9,10,11]
]

my_transform.flatten()
print(my_transform._data_flattened == [0,1,2,3,4,5,6,7,8,9,10,11])
```

Output:
```
True
```

### IMPORTANT!!! At the very bottom of this notebook is a final cell, and you have to write code for it.

# Tests

The cell(s) below contains unit test(s) for your solution.

Run the below cell (after running all the above cells with your code).

<!-- grade_id: cell-dae4d2d23702fe6a, points: 1, grade: true, solution: false -->
```python
assert hasattr(MyTransform, '__init__') and callable(getattr(MyTransform, '__init__')) == True

my_transform = MyTransform()
assert hasattr(my_transform, '_row') == True
assert hasattr(my_transform, '_col') == True
```

<!-- grade_id: cell-935c1f856156a61f, points: 1, grade: true, solution: false -->
```python
assert my_transform.row == 3
assert my_transform.col == 4
```

<!-- grade_id: cell-e96f8241bc0b5aa6, points: 1, grade: true, solution: false -->
```python
my_transform.data_original = [
  [0,1,2,3],
  [4,5,6,7],
  [8,9,10,11]
]
assert hasattr(my_transform, '_data_original') == True
```

<!-- grade_id: cell-ac6764d8bca16895, points: 1, grade: true, solution: false -->
```python
my_transform.flatten()
assert my_transform._data_flattened == [0,1,2,3,4,5,6,7,8,9,10,11]
```

<!-- grade_id: cell-92405c30c42d5f16, points: 1, grade: true, solution: false -->
```python
assert hasattr(my_transform, '_data_flattened') == True
assert hasattr(my_transform, 'data_flattened') == False
```

<!-- grade_id: cell-c109aa2f98a44d2a, points: 1, grade: true, solution: false -->
```python
my_transform.row = 4
my_transform.col = 3

assert my_transform.row == 4
assert my_transform.col == 3

assert my_transform.reshape() == [
  [0,1,2],
  [3,4,5],
  [6,7,8],
  [9,10,11]
]
```

<!-- grade_id: cell-d176c7dce9d591ce, points: 1, grade: true, solution: false -->
```python
my_transform.row = 2
my_transform.col = 6

assert my_transform.row == 2
assert my_transform.col == 6

assert my_transform.reshape() == [
  [0,1,2,3,4,5],
  [6,7,8,9,10,11]
]
```

<!-- grade_id: cell-39f6db96b1d5b4c3, points: 1, grade: true, solution: false -->
```python
assert my_transform._data_flattened == [0,1,2,3,4,5,6,7,8,9,10,11]
```

1 point will be given manually by checking that your code adheres to coding style conventions.

Please add one code cell below, and write code to demonstrate how ``__repr__()`` works.

```python
my_transform = MyTransform()
my_transform.data_original = [
    [0, 1, 2, 3],
    [4, 5, 6, 7],
    [8, 9, 10, 11]
]
my_transform.flatten()

print(my_transform)
```

Output:
```
MyTransform(data_original=[[0, 1, 2, 3], [4, 5, 6, 7], [8, 9, 10, 11]], data_flattened=[0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11], row=3, col=4)
```

```python

```