---

# Assignment 3 Task 1 (1 point)

Write a function named ``multiply`` that multiplies an arbitrarily long list of number arguments (not passed as a list/collection but as single arguments) and returns the result.

<!-- grade_id: cell-810457514e42bdba, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def multiply(*args):
    result=1
    for number in args:
        result*=number
    return(result)
```

# Tests

The cell(s) below contains unit test(s) for your solution.

Run the below cell (after running all the above cells with your code).

<!-- grade_id: cell-336ad2bf9b243033, points: 1, grade: true, solution: false -->
```python
assert multiply(1, 2, 4) == 8
```

```python

```

```python

```