---

# Assignment 4 Task 2 (3 points)

In the same ``mytransform`` module, write a function called ``reshape`` that does the following:
- Input named data: a list of integers
- Input named row: number of rows in the return matrix
- Input named col: number of columns in the return matrix
- Return value: a matrix, as a list of lists of integers, that is the re-arrangement of the input list integers.
- If the data size does not match row and col, raise an appropriate error / exception. Hint: Pick a suitable error from the complete list at the bottom of this webpage: https://docs.python.org/3/library/exceptions.html

- Example 1:
  - Inputs = ``data = [0, 1, 2, 3, 4, 5], row = 2, col = 3``
  - Return value =
    ```
    [
      [0,1,2],
      [3,4,5]
    ]
    ```

- Example 2:
  - Inputs = ``data = [0, 1, 2, 3, 4, 5, 6, 7], row = 4, col = 2``
  - Return value =
    ```
    [
      [0,1],
      [2,3],
      [4,5],
      [6,7]
    ]
    ```

- Example 3:
  - Inputs = ``data = [0, 1, 2, 3, 4], row = 2, col = 2`` --> an error is raised

Write code to import your module, so that the automated test code below can call your funciton like this:

```
reshape(...)
```

<!-- grade_id: cell-6a7b7017d6dc930b, points: 0, grade: false, solution: true -->
```python
from mytransform import reshape

```

### IMPORTANT!!! At the very bottom of this notebook is a final cell, and you have to write code for it.

# Tests

The cell(s) below contains unit test(s) for your solution.

Run the below cell (after running all the above cells with your code).

<!-- grade_id: cell-e2d4dff761c16275, points: 1, grade: true, solution: false -->
```python
assert reshape(data = [0, 1, 2, 3, 4, 5], row = 2, col = 3) == [
  [0,1,2],
  [3,4,5]
]
```

<!-- grade_id: cell-d105309723ef1379, points: 1, grade: true, solution: false -->
```python
assert reshape(data = [0, 1, 2, 3, 4, 5, 6, 7], row = 4, col = 2) == [
  [0,1],
  [2,3],
  [4,5],
  [6,7]
]
```

Please add one code cell below, and write code to demonstrate that an error will be raised when data size does not match row and col. This will be manually graded.

```python
try:
    assert reshape(data=[0, 1, 2, 3, 4, 5, 6], row=4, col=2) == [
        [0, 1],
        [2, 3],
        [4, 5],
        [6, 7]
    ]
except ValueError as e:
    print(e)
```

Output:
```
The total number of elements does not match the desired shape.
```

```python

```