---

# Assignment 3 Task 3 (1 point)

Write a function ```is_sorted``` that takes a list as an argument and checks whether the list is sorted in ascending order or not. The print statements below should all return the expected output. Also write a docstring in the functions head - between the triple quotes - describing the input and the output of the function.

```
    print(is_sorted([-3, 0.5, 100])) --> True
    print(is_sorted([2, 3, 5, 4.5])) --> False
    print(is_sorted(['a', 'c', 'e'])) --> True
    print(is_sorted(['a', 'e', 'c'])) --> False
```

<!-- grade_id: cell-9552dd6d8eb8b442, points: 0, grade: false, solution: true -->
```python
def is_sorted(list_to_sort):
   for i in range(len(list_to_sort)-1): 
        if list_to_sort[i] > list_to_sort[i + 1]:
            return False 
   return True 
```

# Tests

The cell below contains unit tests for your solution.

Run the below cell (after running all the above cells with your code).

<!-- grade_id: cell-73c1b287466cff38, points: 1, grade: true, solution: false -->
```python
assert is_sorted([-3, 0.5, 100]) == True
assert is_sorted(['a', 'e', 'c']) == False
```

```python

```