---

# Assignment 4 Task 1 (3 points)

In ``assignment4_task1_data.txt`` are some data as integers. Write code to read the data from this file, and store the integers in a list called ``data``.

<!-- grade_id: cell-615539137c249d18, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
data=[]
with open("assignment4_task1_data.txt", "r" ) as file:
        for i in file:
            data.append(int(i.strip()))

print(data)

```

Output:
```
[3, 1, 0, 2, 2, 4, 1, 0, 3, 1, 4, 1, 2, 0, 3]
```

For the rest of this assignment, the goal is to implement some common __data transformations__ yourself withOUT using libraries like numpy and pandas. (Calling such functions in a library will lead to zero points!)

Create a module called ``mytransform``, i.e. create a new file ``mytransform.py``.

In this module, write a function called ``onehot`` that does the following:
- Input named data: a list of integers
- Input named classes: number of classes
- Return value: the one-hot encoding (or one-hot vectors) for the input list based on the given number of classes, arranged as a list of lists of integers.
- Example:
  - Inputs = ``data = [3, 0, 1], classes = 5``
  - Return value =
    ```
    [
      [0,0,0,1,0],
      [1,0,0,0,0],
      [0,1,0,0,0]
    ]
    ```
- You can assume that the input data will never violate the number of classes.

See this link for more info about one-hot encoding: https://www.geeksforgeeks.org/ml-one-hot-encoding/

Write code to import your module, so that the automated test code below can call your funciton like this:

```
mt.onehot(...)  # mt is an alias for mytransform
```

<!-- grade_id: cell-f5c4c21a56bf7d62, points: 0, grade: false, solution: true -->
```python
import mytransform as mt

mt.onehot([3,0,1],classes=5)
```

Output:
```
[[0, 0, 0, 1, 0], [1, 0, 0, 0, 0], [0, 1, 0, 0, 0]]
```

# Tests

The cell(s) below contains unit test(s) for your solution.

Run the below cell (after running all the above cells with your code).

<!-- grade_id: cell-c5550f7439e2f16c, points: 1, grade: true, solution: false -->
```python
assert data == [3, 1, 0, 2, 2, 4, 1, 0, 3, 1, 4, 1, 2, 0, 3]
```

<!-- grade_id: cell-36ac2640cec8ffb8, points: 1, grade: true, solution: false -->
```python
assert mt.onehot(data = data, classes = 5) == [
    [0,0,0,1,0],  # 3
    [0,1,0,0,0],  # 1
    [1,0,0,0,0],  # 0
    [0,0,1,0,0],  # 2
    [0,0,1,0,0],  # 2
    [0,0,0,0,1],  # 4
    [0,1,0,0,0],  # 1
    [1,0,0,0,0],  # 0
    [0,0,0,1,0],  # 3
    [0,1,0,0,0],  # 1
    [0,0,0,0,1],  # 4
    [0,1,0,0,0],  # 1
    [0,0,1,0,0],  # 2
    [1,0,0,0,0],  # 0
    [0,0,0,1,0]   # 3
]
```

<!-- grade_id: cell-eaecc53d8533845e, points: 1, grade: true, solution: false -->
```python
assert mt.onehot(data = [0,1,2,3,4,5,6,7,8,9], classes = 10) == [
    [1,0,0,0,0,0,0,0,0,0],
    [0,1,0,0,0,0,0,0,0,0],
    [0,0,1,0,0,0,0,0,0,0],
    [0,0,0,1,0,0,0,0,0,0],
    [0,0,0,0,1,0,0,0,0,0],
    [0,0,0,0,0,1,0,0,0,0],
    [0,0,0,0,0,0,1,0,0,0],
    [0,0,0,0,0,0,0,1,0,0],
    [0,0,0,0,0,0,0,0,1,0],
    [0,0,0,0,0,0,0,0,0,1]
]
```

```python
mt.onehot(data = [0,1,2,3,4,5,6,7,8,9], classes = 10)
```

Output:
```
[[1, 0, 0, 0, 0, 0, 0, 0, 0, 0],
 [0, 1, 0, 0, 0, 0, 0, 0, 0, 0],
 [0, 0, 1, 0, 0, 0, 0, 0, 0, 0],
 [0, 0, 0, 1, 0, 0, 0, 0, 0, 0],
 [0, 0, 0, 0, 1, 0, 0, 0, 0, 0],
 [0, 0, 0, 0, 0, 1, 0, 0, 0, 0],
 [0, 0, 0, 0, 0, 0, 1, 0, 0, 0],
 [0, 0, 0, 0, 0, 0, 0, 1, 0, 0],
 [0, 0, 0, 0, 0, 0, 0, 0, 1, 0],
 [0, 0, 0, 0, 0, 0, 0, 0, 0, 1]]
```

```python

```